<subject>
Assignment How-To
</subject>

<details>
    
**Do Not Talk During Quizzes or Exams**

Do not talk at all once an exam has begun. You may talk again once you leave the room after the quiz or exam. The one exception to this rule is if you need to ask the exam proctor a question. Any talking during a quiz or exam qualifies the student for disciplinary action.

**Naming Conventions**

When naming your files for upload, you must follow the format below:

    <uni>_<assignment>_<details [optional]>.<extension>

For example, if I needed to hand in HW 0, any of the below formats would be sufficient for a file name:

- pl2648_hw0.ipynb
- pl2648_hw0.txt
- pl2648_hw0.sh
- pl2648_hw0_all_in_one.txt
- Pl2648_hw0_bash_program.sh
- Pl2648_quiz1.ipynb

This naming format allows for autograding of all assignments. If your files are not named with this format, you should expect a grade of zero for the assignment.

Courseworks may rename your file to something like `pl2648_hw0-1.ipynb` if you resubmit your assignment. This is perfectly fine.

**What Format To Submit In**

Most homework and quizzes are in Jupyter notebooks. Unless specified otherwise, please download your work as an `.ipynb` file from your local machine and upload it to courseworks.


**Grading**

Possible points on late assignments are deducted by 50% for each day they are late. For example, if you get 80% of the total possible credits on a homework but hand in that homework a day late, you would get 40%. Assignments two days late get zero points.

Once solutions are posted and graded assignments are handed back, students have 1 week to bring their grading discrepancies to a CA for consideration of possible grading errors. 

Because grading is automated, please delete (or comment out) the `raise NotImplmeneted` code before attempting a problem. 

Empty un-editable cells in an assignment are there for a purpose. They will be filled with tests by the automatic grader. Please do not attempt to remove them.

**Getting Help**

Asking for help is a great way to increase your chance of success. However there are some rules. When asking for help (especially from a fellow student), *you can show your helper your code but you can not view theirs*. You work needs to be your own. You can not post screenshots of your current work to Piazza or other tools used for getting help.

If you need to reach out to a CA for help, please do so via Piazza and not via email. Answers given via Piazza will help you as well as other students. Thus, emails will always have a lower priority for response than Piazza questions. If you do email the CA, please make a note of what section you are in. This helps us identify you in courseworks faster. 

Finally, if you do not get a repsonse from a CA within 48 hours, you may email the professor.

**Multiple Choice**

If the question is multiple choice, you will be given several options to choose from and your function will need to return **one** of those options **verbatim** as a string.

For example: 

Which of the following animals bark?

- dogs
- cats
- fish
- trees

A correct answer would be structured in the following way:

```python
def question_animals():
    return 'dogs'
```

You answer will be stripped of left and right white space and lowercased before comparison to the correct answer during grading.
</details>

**Q**: Create a object that represents a store. Have the initialization method take the square footage of the store. Create a method that can calculate the electrical costs for the store given:

- $0.15 per kWh
- kWh per square foot per hour of 0.19

Ensure your class has the following API:

Instance attributes: 
- `square_footage`

Class level attributes:
- DOLLARS_PER_KWH
- KWH_PER_SQUARE_FOOTAGE_PER_HOUR

The method that calculates the electrical cost for the store should be named `electrical_cost`.

[2 points]

<!-- grade_id: cell-e0ff639912440f53, points: 0, grade: false, solution: true -->
```python
class Store:
    DOLLARS_PER_KWH = 0.15
    KWH_PER_SQUARE_FOOTAGE_PER_HOUR = 0.19

    def __init__(self, square_footage):
        self.square_footage = square_footage
    
    def electrical_cost(self, hours):
        kWh = self.square_footage * self.KWH_PER_SQUARE_FOOTAGE_PER_HOUR
        return kWh * self.DOLLARS_PER_KWH * hours
```

<!-- grade_id: cell-e4863bba06d08e3f, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert Store(2000).square_footage == 2000
assert not hasattr(Store, 'square_footage')
### END TESTS
```

<!-- grade_id: cell-47704c4e3e538b07, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert hasattr(Store, 'DOLLARS_PER_KWH')
assert hasattr(Store, 'KWH_PER_SQUARE_FOOTAGE_PER_HOUR')
### END TESTS
```

<!-- grade_id: cell-e9ab2e3e055a711f, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
assert Store(2000).electrical_cost(900) == 51300
### END TESTS
```

**Q**: Create a counter class that takes a string and can produce (via a method named count) a `dict` containing each word in the string as its keys and the counts of each word in the string as its values. Words are any token separated by a space. The class should not be case sensitive (ie. "It" and "it" are the same word).


```python
>>> c = Counter('This is my best string. It is a great string.')
>>> c.count()
{'a': 1,
 'best': 1,
 'great': 1,
 'is': 2,
 'it': 1,
 'my': 1,
 'string.': 2,
 'this': 1}
 ```
 
**Class API**

Instance attributes:
  - "string"
  
You can not use the collections module.
 
[2 points]

<!-- grade_id: cell-0be5d346f805f3fb, points: 0, grade: false, solution: true -->
```python
class Counter:
    def __init__(self, string):
        self.string = string.lower()
    
    def count(self):
        d = {}
        for word in self.string.split():
            if word not in d:
                d[word] = 0
            d[word] += 1
        return d
```

<!-- grade_id: cell-59639940597f3439, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
counter = Counter('This is my best string. It is a great string.')
assert len(counter.count()) == 8
### END TESTS
```

<!-- grade_id: cell-ea0b8e76389a0ae2, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
counter = Counter('This is my best string. It is a great string.')
assert hasattr(counter, 'string')
assert not hasattr(Counter, 'string')
### END TESTS
```

<!-- grade_id: cell-91fa2691c9bea2eb, points: 1, grade: true, solution: false -->
```python

```

**Q**: Update the counter class to have the appropriate methods for presentation when being passed to `repr` and `str`. `repr` should return the object class name and a truncated version of the counter's string. The string should be truncated to 10 characters and an ellipsis added if the string is over 10 characters. `str` should return the object class name and the key count.

You may need to search the Python documentation for the correct methods to implement for this problem.

```python
>>> obj = Counter('Fe Fi Fo Fum, where is my drum')
>>> repr(obj)
<Counter 'Fe Fi Fo F...'>
>>> str(obj)
'Counter(8)'
```

[1 point]

<!-- grade_id: cell-eff92fe808b528d9, points: 0, grade: false, solution: true -->
```python
class Counter:
    def __init__(self, string):
        self.string = string
    
    def count(self):
        d = {}
        for word in self.string.split():
            if word not in d:
                d[word] = 0
            d[word] += 1
        return d

    def __repr__(self):
        string = self.string
        if len(string) > 10:
            string = string[:10] + '...'
            
        return f"<{self.__class__.__name__} '{string}'>"
    
    def __str__(self):
        return f'{self.__class__.__name__}({len(self.count())})'
```

<!-- grade_id: cell-42a5d9c016cae56f, points: 0.25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-76d7e759c0c50b24, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
counter = Counter('This is my best string. It is a great string.')
assert repr(counter) == "<Counter 'This is my...'>"
### END TESTS
```

<!-- grade_id: cell-a6990dafdb2b143c, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
counter = Counter('Woot woot!')
assert repr(counter) == "<Counter 'Woot woot!'>"
### END TESTS
```

<!-- grade_id: cell-d5d109b50b154b18, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
counter = Counter('This is my best string. It is a great string.')
assert str(counter) == 'Counter(8)'
### END TESTS
```

**Q**: Create a custom exception for a severe lack of parties. It should be called `SevereLackOfPartiesError` and should subclass `ValueError`. The class should have a docstring that describes what the error represents. Use your imagination for the exact text. 

[1 point]

<!-- grade_id: cell-7368367ae633aef9, points: 0, grade: false, solution: true -->
```python
class SevereLackOfPartiesError(ValueError):
    """An utter failure in getting the party started."""
    pass
```

<!-- grade_id: cell-78f404be90d25d2e, points: 0.5, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-1b63d2e6166363a1, points: 0.5, grade: true, solution: false -->
```python

```

**Q**: Create a class `Str` that subclasses `str`. Add a method to the subclass that checks if the string does not start with a given string. Call that method `notstartswith`. This class should only take up 3 lines of code however you will not be graded on the length of your class.

[1 point]

<!-- grade_id: cell-931411951b8973e6, points: 0, grade: false, solution: true -->
```python
class Str(str):
    def notstartswith(self, s):
        return not self.startswith(s)
```

<!-- grade_id: cell-dde9e6cf99caf6bc, points: 0.5, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-767d16e4e70c6601, points: 0.25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-7f7773df3332732e, points: 0.25, grade: true, solution: false -->
```python

```

**Q**: Create a `BankAccount` class that is instantiated with an account number and an optional initial balance. Add methods to withdraw and deposit to the account's balance. Also, add a method that allows for the casting of objects of this class to a string.

```python
a = BankAccount(12345678, balance=10)
a.deposit(5)
a.withdraw(3)
str(a)

'Account ID: 12345678 - Balance: 12'
```

[2 points]

<!-- grade_id: cell-31e546afd67a9412, points: 0, grade: false, solution: true -->
```python
class BankAccount:
    def __init__(self, account_id, balance=0):
        self.account_id = account_id
        self.balance = balance

    def withdraw(self, amount):
        self.balance -= amount
        return self.balance

    def deposit(self, amount):
        self.balance += amount
        return self.balance
    
    def __str__(self):
        return f'Account ID: {self.account_id} - Balance: {self.balance}'
```

<!-- grade_id: cell-78f78a2ea5d08b99, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert str(BankAccount(4839283, balance=2010)) == 'Account ID: 4839283 - Balance: 2010'
### END TESTS
```

<!-- grade_id: cell-e706bb44eae7db75, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert BankAccount(1).balance == 0
### END TESTS
```

<!-- grade_id: cell-da11f9aeadad2c39, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
a = BankAccount(1)
a.deposit(150)
assert a.balance == 150
### END TESTS
```

<!-- grade_id: cell-0f6bc2f6bd8c7a9c, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
a = BankAccount(1)
a.deposit(150)
a.withdraw(27)
assert a.balance == 123
### END TESTS
```

**Q**: Subclass the `BankAccount` class from the previous question to create a `MinBalanceBankAccount` class. This class should require a minimum balance at all times and should raise an exception (of the type `ValueError`) when an attempt to drop the balance below the minumum occurs. The `MinBalanceBankAccount` should be instantiated with the account number, inital balance, and the minimum balance. No error checking against the minimum balance is necessary upon instantiation.

```python
a = MinBalanceBankAccount(account_id=1, balance=100, minimum_balance=10)
a.withdraw(100)
...
ValueError: Sorry, minimum balance must be maintained.
```

If raised, the `ValueError` should have the message: 'Sorry, minimum balance must be maintained.'

A minimum of extra code should be implemented in `withdraw` to meet the requirements of this question.

[2 points]

<!-- grade_id: cell-ffede70c71976fef, points: 0, grade: false, solution: true -->
```python
class MinBalanceBankAccount(BankAccount):
    def __init__(self, account_id, balance, minimum_balance):
        self.minimum_balance = minimum_balance
        super().__init__(account_id, balance)
        
    def withdraw(self, amount):
        if self.balance - amount < self.minimum_balance:
            raise ValueError('Sorry, minimum balance must be maintained.')

        super().withdraw(amount)
```

<!-- grade_id: cell-4d189116d6cb3d51, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
try:
    MinBalanceBankAccount(1, 100, 10).withdraw(100)
    assert False, 'Should have raised value Error'
except ValueError as e:
    assert str(e) == 'Sorry, minimum balance must be maintained.'
### END TESTS
```

<!-- grade_id: cell-496de9ecb2b38244, points: 0.5, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-74b6dba911c7df84, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert issubclass(MinBalanceBankAccount, BankAccount)
### END TESTS
```

<!-- grade_id: cell-3d7deeb06f900ce3, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
a = MinBalanceBankAccount(1, 100, 10)
a.deposit(150)
a.withdraw(27)
assert a.balance == 223
### END TESTS
```

**Extra Practice** Create a class that produces objects that are both iterables and iterators. 

<!-- grade_id: cell-acfb6cbf62089d2f, points: 0, grade: false, solution: true -->
```python
class MyRange:
    def __init__(self, start, stop):
        self.start = 0
        self.stop = stop
        self.pos = None

    def __iter__(self):
        return self  
        # Hand back the object that will keep track 
        # of progress as we step through iterable.
    
    def __next__(self):
        if self.pos is None:
            self.pos = self.start
        else:
            self.pos += 1
        
        if self.pos >= self.stop:
            raise StopIteration

        return self.pos
# YOUR CODE HERE
```

<!-- grade_id: cell-ef5f590e0182d168, points: 0, grade: true, solution: false -->
```python
### BEGIN TESTS
assert list(MyRange(1, 5)) == [0, 1, 2, 3, 4]
### END TESTS
```

```python

```